## DESCRIPTION
## Calculus
## ENDDESCRIPTION


## Tagged by tda2d

## DBsubject(Calculus - multivariable)
## DBchapter(Integration of multivariable functions)
## DBsection(Change of variable)
## Date(01/30/2011)
## Institution(Hope College)
## Author(Paul Pearson)
## Level(3)
## KEYWORDS('Multivariable','change of variables')


## SOURCE 
## /RELCALC3/set08/Hope-Multi2_11-04-MV-Int-change-var_Change-variables-05.pg' 

###########################
#  Initialization

DOCUMENT();

loadMacros(
  "PGstandard.pl",
  "MathObjects.pl",
  "parserMultiAnswer.pl",
  "PGmatrixmacros.pl",
  "PGunion.pl",
  "PGcourse.pl"
);

TEXT(beginproblem());

SRAND($psvn);


###########################
#  Setup

Context("Numeric")->variables->are(u=>'Real',v=>'Real');

$a = random(2,4,1);
$b = random(5,7,1);

$aa = $a * $a;
$bb = $b * $b;

$x = Compute("u^2 / v");
$y = Compute("v / u");

$ulo = Real(1);
$uhi = Real($aa);
$vlo = Real(1);
$vhi = Real($bb);


$multians0 = MultiAnswer( $x, $y )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>  </table>",
  tex_format =>  "\begin{array}{r} %s \\ %s \end{array}",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..1) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$multians1 = MultiAnswer( Compute("2 u / v"), Compute("-u^2 / v^2"), Compute("-v / u^2"), Compute("1 / u"), Compute("1 / v") )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='10'>
  <tr><td> %s </td><td> %s </td></tr>
  <tr><td> %s </td><td> %s </td></tr>
  </table> %s",
  tex_format =>  "\left\lbrack\begin{array}{rr} %s & %s \\ %s & %s \end{array}\right\rbrack; %s",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..4) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);


$area = Compute("($aa - 1) * ln($bb)");







###########################
#  Main text

Context()->texStrings;
BEGIN_TEXT
Calculer l'aire de la région \( D \) bornée par
$PAR
\[
xy = 1, \
xy = $aa, \ \ \
xy^2 = 1, \ 
xy^2 = $bb
\]
$BR
dans le premier quadrant du plan \( xy \)

$BR
$BR
(a) Dessiner un graphe représentant la région \( D \).
$BR
$BR
(b) Utiliser le changement de variable non-linéaire

\( u = xy \) et \( v = xy^2 \), trouver \( x \)
et \( y \) comme fonction de \( u \) et \( v \).
$BR
$BR
\( x = x(u,v) = \) \{ $multians0->ans_rule(20) \}
$BR
\( y = y(u,v) = \) \{ $multians0->ans_rule(20) \}
$BR
$BR
(c)  Calculer le déterminant du Jacobien de ce changement de variable. 

$PAR
$BCENTER
\{ 
mbox(
  "\( \displaystyle \left| \frac{\partial(x,y)}{\partial(u,v)} \right| = \mathrm{det} \)",
  display_matrix([
  [$multians1->ans_rule(10),$multians1->ans_rule(10)], 
  [$multians1->ans_rule(10),$multians1->ans_rule(10)]], align=>'cc'),
  " = ".$SPACE.$multians1->ans_rule(20)
);
\}
$ECENTER
$BR
(d) Utiliser ce changement de variable calculer l'aire de la région \( D \) avec une intégrale double.
$BR
$BR
END_TEXT
Context()->normalStrings;

############################
#  Answer evaluation

$showPartialCorrectAnswers = 1;

#ANS( $x->cmp() );
#ANS( $y->cmp() );

ANS( $multians0->cmp() );

ANS( $multians1->cmp() );








$multians2 = MultiAnswer( Compute("$bb"), Compute("1"), Compute("$aa"), Compute("1"), Compute("1/v") )->with(

  singleResult => 1,
  allowBlankAnswers => 1,
  checkTypes => 0,
  format =>  "<table border='0' cellspacing='0'><tr><td>
  <table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table></td> 
  <td><table border='0' cellspacing='10'>
  <tr><td> %s </td></tr>
  <tr><td> %s </td></tr>
  </table></td>
  <td>%s</td></tr></table>",
  tex_format =>  "\int^{%s}_{%s} \int^{%s}_{%s} %s du dv",

  checker => sub {
      my ( $correct, $student, $answerHash ) = @_;
      my @c = @{$correct};
      my @s = @{$student};
      my @score = ();
      foreach my $j (0..4) {
        $score[$j] = 0;
        if ($c[$j]->typeMatch($s[$j]) && $c[$j] == $s[$j]) { $score[$j] = 1; }
      }
      return [ @score ];
  }

);




#$integraltex = "\int_a^b \int_c^d f(x(u,v),y(u,v)) \left| \frac{\partial(x,y)}{\partial(u,v)} \right| du \, dv";

$integraltex = "\int\!\int_D dx \, dy = \int_a^b \int_c^d \left| \frac{\partial(x,y)}{\partial(u,v)} \right| du \, dv";

#
#  Display the answer blanks properly in different modes
#
Context()->texStrings;
if ($displayMode eq 'TeX') {
   $integral =
   "\( \displaystyle $integraltex = \) \\\\ \( \displaystyle ".
   '\int_{'.$multians2->ans_rule(4).'}^{'.$multians2->ans_rule(4).'}'. 
   '\int_{'.$multians2->ans_rule(4).'}^{'.$multians2->ans_rule(4).'}'. 
   $multians2->ans_rule(20).'du \, dv \)';
  } else {
   $integral =
   BeginTable(center=>0).
     Row([
       "\( \displaystyle $integraltex = \int \)",
       $multians2->ans_rule(4).$BR.$BR.$multians2->ans_rule(4),
       '\( \displaystyle \int \)',
       $multians2->ans_rule(4).$BR.$BR.$multians2->ans_rule(4),
       $multians2->ans_rule(20), '\( du \, dv \)'],separation=>2).
   EndTable();
}
Context()->normalStrings;



Context()->texStrings;
BEGIN_TEXT
$integral
$BR
$BR
(e) Évaluer l'intégrale double pour trouver la valeur de l'aire de la région \( D\).
$BR
$BR
Aire = \{ ans_rule(20) \} 
END_TEXT
Context()->normalStrings;


#ANS( Compute("1")->cmp() ); # A 
#ANS( Compute("$aa")->cmp() ); # B
#ANS( Compute("1")->cmp() ); # C
#ANS( Compute("$bb")->cmp() ); # D
#ANS( Compute("1 / v")->cmp() ); # Integrand * Jacobian

ANS( $multians2->cmp() );

ANS( $area->cmp() );



ENDDOCUMENT();
